Food-washing monkeys recognize the law of diminishing returns

Author response:

The following is the authors’ response to the previous reviews

We thank the editors and Reviewers 1 and 3 for their though6ul consideration of our manuscript. The present revision is submitted to address comments raised concerning rank determinations and the following sentence in the editorial assessment:

The evidence that food-washing is deliberate is compelling, but the evidence for variable and adaptive investment depending on rank, including the fitness-relevance and ultimate evolutionary implications of the findings, is incomplete given limitations of the experimental design.

Close reading of this sentence reveals two parallel threads. The first can be read as “…evidence for variable rank is incomplete given the limitations of the experimental design,” whereas the second can be read as “…evidence for adaptive investment and fitness is incomplete given the limitations of the experimental design.” The first alludes to a critique of our methods, while the second alludes to points of discussion unrelated to our experimental design. Unpacking this sentence is important because it casts the totality of our paper as “incomplete,” a word of consequence for early-career scholars because it prevents indexing in Web of Science.

For clarity, we will refer to these topics as Thread 1 and Thread 2 in the following response.

Thread 1 seems rooted in a comment made by Reviewer 1, which is reproduced below:

I am still struck that there was an analysis of only trials where <3 individuals are present. If rank was important, I would imagine that behavior might be different in social contexts when theA, scrounging, policing, aggression, or other distractions might occur-- where rank would have effects on foraging behavior. Maybe lower rankers prioritize rapid food intake then. If rank should be related to investment in this behavior, we might expect this to be magnified (or different) in social contexts where it would affect foraging. It might just be that the data was too hard to score or process in those settings, or the analysis was limited. Additionally, I think that more robust metrics of rank from more densely sampled focal follow data would be a beJer measure, but I acknowledge the limitations in getting the ideal. Since rank is central to the interpretation of these results, I think that reduced social contexts in which rank was analyzed and the robustness of the data from which rank was calculated and analyzed are the main weaknesses of the evidence presented in this paper.

We are grateful for this perspective of Reviewer 1, but it puts us in an uncomfortable position. We must respond rather forcefully because of its influence on the above assessment. A problem with R1’s comment is that it uses the word “foraging” (a behavior we did not study) instead of “cleaning” (the behavior we did study). Still, we can substitute the latter word with the former to get the gist of it.

R1 criticizes our methods as a prelude for imagining the behaviors of our study animals, a form of conjecture. R1 correctly supposes a positive relationship between the number of animals and the intensity of competition for a limited food resource, a well-known phenomenon; and, yes, the food in each trial was decidedly limited, being fixed at nine cucumber slices. But R1 incorrectly presumes rank effects on cleaning under conditions of intense food competition. When the number of monkeys participating in a trial exceeded the number of feeding stations (n = 3), we saw little or no cleaning effort, either brushing or washing. So, rank effects on cleaning are immaterial under these conditions. As our study goals were narrowly focused on detecting individual propensities, or choices, as a function of rank, we limited our analysis to trials involving three monkeys or fewer. In retrospect, we admit that we should have provided better justification for our choice of trials, so we’ve edited one of our sentences:

Original sentence

Formerly lines 219-220: To minimize the potential confounding effects of dominance interactions, we analyzed trials with ≤ 3 monkeys.

Revised sentence

Current lines 219-224: We excluded trials from analysis if the number of participating monkeys exceeded the number of feeding stations, as these conditions produced high levels of feeding competition with scant cleaning behavior. Such conditions effectively erased individual variation in sand removal, the topic motivating our experiment. Accordingly, we analyzed trials with ≤ 3 monkeys, putting 937 food-handling bouts into the GLMM statistical models, which included data on individual rank, sex, and sand treatment.

R1’s final criticism – “I think that more robust metrics of rank from more densely sampled focal follow data would be a better measure, but I acknowledge the limitations in getting the ideal” – seems to imply that rank data were collected during our experiment. On the contrary, we determined ranks from five years of focal follows preceding the experiment, achieving the very standard that R1 describes as ideal. The relevant text appeared on lines 165-169 in version 2.0:

To determine the rank-order of adults, we recorded dyadic agonistic interactions and their outcomes (i.e., aggression, supplants, and silent-bared-teeth displays of submission) during 5min focal follows of individuals based on a randomized order of continuous rotation (Tan et al., 2018). In some cases, these data were supplemented with ad libitum observations. This protocol existed during five years (2013-2018) of continual observations before we conducted our experiment in July-August 2018.

Naturally, we were puzzled by R1’s dismissal of our methods, as well as R1’s conclusion, reached without evidence, that “[the] reduced social contexts in which rank was analyzed and the robustness of the data from which rank was calculated and analyzed are the main weaknesses of the evidence presented in this paper.” It is unsubstantiated assertation with no definition of robustness, making it difficult for anyone to objectively assess the quality of our data.

We detect in R1’s words some unfamiliarity with the social organization of our study species, which is fair enough. To better orient readers to the dominance hierarchy of Macaca fascicularis, and to boost reader confidence in the volume and quality of our rank data, we have added several sentences to this section of the manuscript, lines 169-183:

Macaques form multi-male multi-female (polygynandrous) social groups with individual dominance hierarchies. In M. fascicularis, the hierarchy is strictly linear and extremely steep, meaning aggression is unidirectional (de Waal, 1977; van Noordwijk and van Schaik, 2001) with profound asymmetries in outcomes for individuals of adjacent ranks (Balasubramaniam et al., 2012). Further, the dominance hierarchies of philopatric females are stable and predictable. Daughters follow the pattern of youngest ascendancy, ranking just below their mothers with few known exceptions among older sisters (de Waal, 1977; van Noordwijk and van Schaik, 1999). Taken together, these species traits are conducive to unequivocal rank determinations.

To determine the rank-order of adults in our study group, we recorded dyadic agonistic interactions and their outcomes (i.e., aggression, supplants, and silent-bared-teeth displays of submission) during 5-min focal follows of individuals based on a randomized order of continuous rotation (Tan et al., 2018). These data were supplemented with ad libitum observations and all rank determinations were updated monthly, and when males immigrated or emigrated. This protocol predates our experiment in July-August 2018, representing 970 hr of focal data during five years of systematic study (2013-2018).

Thread 2 criticizes our evidence for adaptive investment and fitness, describing it is a limitation of our experimental design. Accordingly, the totality of our experiment was classified as “incomplete.” Yet, our experiment was never designed to collect such evidence, and we make no claims of having it. Rather, we discussed potential fitness consequences to highlight the broader significance of our study, connecting it diverse bodies of literature, from evolutionary theory to paleoanthropology. Our intent was to follow the conventions of scientific writing; to put our results into conversation with the wider literature and set an agenda for future research.

On reflection, Thread 2 seems to pivot around something as arbitrary as structure. Previously, our results and discussion were combined under a single section header (“Results and Discussion”), a stylistic choice to economize words. Our manuscript is a Short Report, which is limited to 1,500 words of main text. But this level of concision proved counterproductive. It blurred our results and discussion in the minds of readers. Indeed, Reviewer 3 described it as “misleading,” a barbed word that accomplishes the same act attributed to us. To counter this perspective, we have simply partitioned our Results (now “Experimental Results”) and Discussion to draw a sharper distinction between the two components of our paper.

Changes from the earlier version to the revised version

Abstract

Abstract

Few animals have the cognitive faculties or prehensile abilities needed to eliminate tooth-damaging grit from food surfaces. Some populations of monkeys wash sand from foods when standing water is readily accessible, but this propensity varies within groups for reasons unknown. Spontaneous food-washing emerged recently in a group of long-tailed macaques (Macaca fascicularis) inhabiting Koram Island, Thailand, and it motivated us to explore the factors that drive individual variability. We measured the mineral and physical properties of contaminant sands and conducted a field experiment, eliciting 1,282 food-handling bouts by 42 monkeys. Our results verify two long-standing presumptions, that monkeys have a strong aversion to sand and that removing it is intentional. Reinforcing this result, we found that monkeys clean foods beyond the point of diminishing returns, a suboptimal behavior that varied with social rank. Dominant monkeys abstained from washing, a choice consistent with the impulses of dominant monkeys elsewhere: to prioritize rapid food intake and greater reproductive fitness over the long-term benefits of prolonging tooth function.

Introduction

Koshima Island, Japan is a storied fieldsite in the annals of primatology. Observations of wild macaques (Macaca fuscata) began in 1948, but four years of sustained effort failed to habituate the monkeys. In August 1952, Itani and Tokuda (1954) resorted to scattering wheat grains and sweet potatoes along oft-used paths, gradually shifting the provisions to a sandy beach at Otomari Bay, a strategy that afforded a clear view of the entire group of 22 animals. Individual identifications followed quickly, laying the foundation for decades of influential research. In September 1953, a young female named Imo gathered a sweet potato from the beach and rinsed it in a freshwater stream, a behavioral innovation that spread horizontally to peers and then vertically to older kin (Kawai, 1965). By 1958, sweet potato washing had become a group-wide trait with a key modification: the monkeys began using sea water instead of standing freshwater, a preference that continues today seven generations later (Hirata et al., 2001).

These events have since passed into canon as an example of socially transmitted behavior, or culture, among nonhuman primates (McGrew, 1998; Matsuzawa and McGrew, 2008; Matsuzawa, 2015). It is a venerable legacy, but it overshadows a fundamental question: why do monkeys wash their food? (Sarabian and MacIntosh, 2015; Fiore et al., 2020). Two tacit assumptions—that sand produces an objectionable sensation on teeth, and that it is prudent to minimize tooth damage—are sufficiently intuitive that formal tests are wanting. In consequence, the mineral and physical properties of contaminant sands are unknown, let alone the efficiency of different cleaning behaviors (Schofield et al., 2018). Another enigma concerns the preference of some monkeys to brush food with their hands (Kawai et al., 1992), a rapid but seemingly inferior means of sand removal. Foodbrushing individuals have been characterized as inept (Kawai, 1965) or subordinate (Watanabe, 1994) in part because the quartz in sand can cause severe tooth damage (Lucas et al., 2013; Towle et al., 2022). Yet, carrying food to the ocean is expected to incur energetic costs and opportunity costs, factors that impelled us to explore the trade-offs of mitigating sand-mediated tooth wear.

Koram Island, Thailand

The long-tailed macaques (M. fascicularis) of Koram Island, Thailand use stone tools to harvest shellfish, a phenomenon that came to light during the surveys of biodiversity damage that followed the Sumatra-Andaman earthquake and tsunami of December 26, 2004 (Malaivijitnond et al., 2007; Gumert and Malaivijitnond, 2012; Tan et al., 2015). The monkeys then became a magnet for tourism. Some visitors began supplying the monkeys with market-sourced fruits (cucumbers, melon, pineapple), jettisoning them onto the beach. These events produce food surfaces with considerable concentrations of sand ( x¯ = 3.7 ± 1.3 mg mm−2), which, in turn, elicit food-washing and food-brushing behaviors among the monkeys. To understand the factors driving these reactions, we examined the mineral properties of food-adhering sands (n = 758 particles), finding that 78% of our sample was composed of crystalline quartz (Figure 1A).

Variation in the mineral and physical properties of quartz particles on food surfaces.

(A) Particle sizes followed a bimodal distribution, with most particles featuring metal inclusions. Nearly half the sample is < 25 μm, the “grittiness threshold” of the human oral cavity (Imai et al., 1995). (B) Circularity is a dimensionless shape factor (range: 0 to 1) based on two-dimensional microscopy and estimates for the projected area and perimeter of a given particle. Some examples are illustrated; overall, it is a convenient but imperfect proxy for sphericity, or deviations from spherical (Grace and Ebneyamini, 2021). Here, circularity varied as a function of mean Feret diameter, suggesting that larger particles hold greater potential for damaging attack angles during particle-enamel contact.

Harder than enamel, quartz can exact a heavy toll on teeth, but the probability and degree of enamel loss is governed partly by the size and shape of individual particles, factors that determine the ‘attack angle’ during particle-enamel contact (Lucas et al., 2013). We calculated the circularity of particles as a convenient proxy for sphericity (Grace and Ebneyamini, 2021), finding that it decreased as a function of particle size (Figure 1B). This result suggests that larger particles are more angular, posing a greater risk to enamel. However, we also calculated a median Feret diameter of 25.8 μm (range: 8.7 to 644 μm), meaning that nearly half the sample existed below the human threshold (25 μm) of oral detection (Imai et al., 1995). To put 25 μm into perspective, it is one-twelfth the diameter of the period ending this sentence.

Mitigating tooth wear

Chewing undetected quartz is expected to cause severe tooth wear, but this cost can be mitigated behaviorally if food-cleaning is proficient. To test this contention, we simulated the brushing and washing actions of monkeys with cucumber slices exposed to three concentrations of sand: low ( x¯ = 0.2 ± 0.2 mg mm−2), intermediate ( x¯ = 0.9 ± 0.1 mg mm−2), and high ( x¯ = 1.8 ± 0.9 mg mm−2). We found that brushing was less efficient than washing across treatments, eliminating 76 ± 7% vs. 93 ± 4% of sand particles, respectively (Figure S1). It is a modest difference, perhaps, but it is freighted with fitness consequences when extrapolated over years of life (Fannin et al., 2022). It follows that monkeys should compulsively wash sand from food whenever the opportunity avails itself, a prediction that motivated a field experiment.

Added in the revised version: Results

Removed in the revised version: Results and discussion

Earlier version: Our experiment was designed to test two concepts at once (Figure 2; Video 1). The first pivots around intentionality, a thorny problem that emerged from studies of raccoons (Procyon lotor). Celebrated food-handlers, the submersion of food objects in water is better termed ‘dousing’ for greater haptic sensation, not washing with the intention of removing surface contaminants (Box 1). If the intent of monkeys is to eliminate sand, then the time devoted to brushing or washing food should vary as a positive function of sandiness. The other concept draws on observations from Koshima Island, which alluded to rank effects on individual cleaning behaviors, a pattern that is difficult to detect without controlling access to food or distance to the ocean.

Revised version: Our experiment was designed to test two concepts at once. The first pivots around intentionality, a thorny problem that emerged from studies of raccoons (Procyon lotor). Celebrated food-handlers, the submersion of food objects in water is better termed ‘dousing’ for greater haptic sensation, not washing with the intention of removing surface contaminants (Box 1). If the intent of monkeys is to eliminate sand, then the time devoted to brushing or washing food should vary as a positive function of sandiness. The other concept draws on observations from Koshima Island, which alluded to rank effects on individual cleaning behaviors, a pattern that is difficult to detect without controlling access to food or distance to the ocean.

Added in the revised version: Manual prehension and the nomenclature of raccoons

Added in the revised version: Raccoons (Procyon lotor) tend to live in wooded habitats near waterways, where they can be seen dousing foods before ingestion (photograph by Markus Zindl, reproduced with permission).

Added in the revised version: Animal names tend to reflect salient features of their appearance or behavior, including the sounds they make. This basic principle of ethnotaxonomy harmonizes the classification of animals with their traits, a pattern that extends to other languages when names are borrowed or translated. But this process can have profound consequences, shaping our understanding of animals and their behaviors. For example, manipulative actions are baked into the word raccoon, a corruption of the Powhatan words arakun or arakunem, meaning, approximately, ‘it scratches with its hands.’ Far more common across North America is the Anishinaabe word esiban (‘it picks up things’), with cognates in the languages of at least a dozen cultural groups, including Cree, Potawatomi, Abenaki, and Delaware peoples (Holmgren, 1990; Justice, 2021). The Tsimshian word que-o-koo (‘washes with hands’) is notably different for imputing intent, perhaps because coastal peoples are sensitized to sand on their foods.

Added in the revised version: Linnaeus put raccoons in the family Ursidae—bears—in the second edition of Systema Naturae (1740), classifying them as Ursus cauda elongata (‘bear with long tail’). By the 10th edition (1758), he had reclassified them as Ursus lotor (‘washer bear’) due to accounts of captive raccoons persistently dunking food in water, along with his own observations of Sjupp, a pet raccoon gifted to him by the Swedish crown prince Adolf Fredrik (Nicholls, 2007). Twenty years later, in 1780, the German naturalist Gottlieb Conrad Christian Storr elevated raccoons into their own genus, Procyon, meaning’before the dog’or’early dog’, a nod to their dog-like appearance;but he retained the species nomen lotor. Today, the raccoon is known as tvattbjorn in Linnaeus’s Swedish and waschbar in Storr’s German, both meaning ‘washer bear.’

Added in the revised version: Meanwhile, the French naturalist Georges-Louis Leclerc, Comte de Buffon conducted his own detailed observations of pet raccoons held in the Muséum National d’Histoire Naturelle, Paris. He saw affinities with rodents, a legacy that echoes today in the names raton laveur and ratão-lavadeiro (‘washing rat’) in French and Portuguese, respectively. So, whether the raccoon is related to bears, dogs, or rats, every European taxonomist of the 18th century agreed that washing was its defining trait. But it is a popular misconception.

Added in the revised version: Raccoons wet their foods to enhance haptic sensation, not eliminate contaminants;their intent is to douse food objects, not wash them (Lyall-Watson, 1963). Intriguingly, the posterior cerebrum—which includes the somatosensory cortex and several other cortical, thalamic, and subcortical structures—is relatively expanded among raccoons, suggesting magnification of the neural pathways that serve tactile processing (Arsznov and Sakai, 2013). When viewed in this light, the Lenape ethnonym nachenum (‘they use hands as tools’) would seem most fitting of all, not least for blurring the thin line between human and nonhuman curiosity and the nature of knowing.

Added in the revised version: We conducted 101 feeding trials, recording 1,282 food-handling events by 42 individuals (Figure 2; Video 1). We had detailed rank information for 23 individuals, so we used this subset of data in our GLMM analyses. We found that all monkeys were sensitive to sand on their food, responding to each treatment—low, intermediate, and high concentrations—with greater median durations (± 1 SD) of brushing (low: 0.0 ± 0.1 s; intermediate: 1.1 ± 2.0 s; high: 3.1 ± 2.0 s; Figure 2B) and washing (low: 0.04 ± 0.3 s; intermediate: 0.6 ± 2.0 s; high: 3.3 ± 4.3 s; Figure 2C). This result is important for upholding long-held assumptions of intentional cleaning. Further, we found that dominant monkeys of both sexes showed a strong propensity for food-brushing (Figure S2A; Table S3) over food-washing (Figure S2B; Table S4). This finding reverses the pattern observed on Koshima Island (Watanabe, 1994), and it raises the possibility that food-washing is an indulgence subject to diminishing returns. To explore this premise, we developed a theoretical model where the time devoted to food-cleaning is predicted to maximise the rate of sand removal as a function of handling time.

Experimental design and results.

(A) To elicit food-cleaning behaviors, we put sliced cucumbers in trays representing three treatments—food surfaces with low ( x¯ = 0.2 ± 0.2 mg mm−2), intermediate ( x¯ = 0.9 ± 0.1 mg mm−2), and high ( x¯ =1.8 ± 0.9 mg mm−2) concentrations of sand—positioned 1.5 m apart and 15 m from the ocean. (B) Monkeys brushed the sandier treatments for longer durations [χ2 (2, n = 575 food-handling bouts) = 194.7, p < 0.0001] with no effect of dominance rank or sex (Table S1; Figure S2). (C) Monkeys washed the sandier treatments for longer durations [χ2 (2, n = 362 food-handling bouts) = 69.7, p < 0.0001], and we found an interaction effect with rank independent of sex [χ 2 = 19.3, p < 0.0001; Table S2; Figure S2]. (D) Energy intake rates also varied as function dominance rank [ANOVA, (LN-transformed); F2,104 = 10.0; p < 0.0001]. Symbols represent mean values and whiskers ± 1 s.e. Photos by Amanda Tan.

Removed in the revised version: Box 1. Manual prehension and the nomenclature of raccoons

Removed in the revised version: Raccoons (Procyon lotor) tend to live in wooded habitats near waterways, where they can be seen dousing foods before ingestion (photograph by Markus Zindl, reproduced with permission).

Removed in the revised version: Animal names tend to reflect salient features of their appearance or behavior, including the sounds they make. This basic principle of ethnotaxonomy harmonizes the classification of animals with their traits, a pattern that extends to other languages when names are borrowed or translated. But this process can have profound consequences, shaping our understanding of animals and their behaviors. For example, manipulative actions are baked into the word raccoon, a corruption of the Powhatan words arakun or arakunem, meaning, approximately, ‘it scratches with its hands.’ Far more common across North America is the Anishinaabe word esiban (‘it picks up things’), with cognates in the languages of at least a dozen cultural groups, including Cree, Potawatomi, Abenaki, and Delaware peoples (Holmgren, 1990; Justice, 2021). The Tsimshian word que-o-koo (‘washes with hands’) is notably different for imputing intent, perhaps because coastal peoples are sensitized to sand on their foods.

Removed in the revised version: Linnaeus put raccoons in the family Ursidae—bears—in the second edition of Systema Naturae (1740), classifying them as Ursus cauda elongata (‘bear with long tail’). By the 10th edition (1758), he had reclassified them as Ursus lotor (‘washer bear’) due to accounts of captive raccoons persistently dunking food in water, along with his own observations of Sjupp, a pet raccoon gifted to him by the Swedish crown prince Adolf Fredrik (Nicholls, 2007). Twenty years later, in 1780, the German naturalist Gottlieb Conrad Christian Storr elevated raccoons into their own genus, Procyon, meaning’before the dog’or’early dog’, a nod to their dog-like appearance;but he retained the species nomen lotor. Today, the raccoon is known as tvättbjörn in Linnaeus’s Swedish and waschbär in Storr’s German, both meaning ‘washer bear.’

Removed in the revised version: Meanwhile, the French naturalist Georges-Louis Leclerc, Comte de Buffon conducted his own detailed observations of pet raccoons held in the Muséum National d’Histoire Naturelle, Paris. He saw affinities with rodents, a legacy that echoes today in the names raton laveur and ratão-lavadeiro (‘washing rat’) in French and Portuguese, respectively. So, whether the raccoon is related to bears, dogs, or rats, every European taxonomist of the 18th century agreed that washing was its defining trait. But it is a popular misconception.

Removed in the revised version: Raccoons wet their foods to enhance haptic sensation, not eliminate contaminants;their intent is to douse food objects, not wash them (Lyall-Watson, 1963). Intriguingly, the posterior cerebrum—which includes the somatosensory cortex and several other cortical, thalamic, and subcortical structures—is relatively expanded among raccoons, suggesting magnification of the neural pathways that serve tactile processing (Arsznov and Sakai, 2013). When viewed in this light, the Lenape ethnonym nachenum (‘they use hands as tools’) would seem most fitting of all, not least for blurring the thin line between human and nonhuman curiosity and the nature of knowing.

Removed in the revised version: We conducted 101 feeding trials, recording 1,282 food-handling events by 42 individuals. We had detailed rank information for 23 individuals, so we used this subset of data in our GLMM analyses. We found that all monkeys were sensitive to sand on their food, responding to each treatment—low, intermediate, and high concentrations—with greater median durations (± 1 SD) of brushing (low: 0.0 ± 0.1 s; intermediate: 1.1 ± 2.0 s; high: 3.1 ± 2.0 s; Figure 2B) and washing (low: 0.04 ± 0.3 s; intermediate: 0.6 ± 2.0 s; high: 3.3 ± 4.3 s; Figure 2C). This result is important for upholding long-held assumptions of intentional cleaning. Further, we found that dominant monkeys of both sexes showed a strong propensity for food-brushing (Figure S2A; Table S3) over food-washing (Figure S2B; Table S4). This finding reverses the pattern observed on Koshima Island (Watanabe, 1994), and it raises the possibility that food-washing is an indulgence subject to diminishing returns. To explore this premise, we developed a theoretical model where the time devoted to food-cleaning is predicted to maximise the rate of sand removal as a function of handling time.

Figure 3A illustrates the fastidious nature of our study population: monkeys allocated excess time to washing and brushing—by factors of 1.5 and 3.0, respectively—beyond that predicted by the optimization of sand removal (Video 2). Our model also highlights sharply divergent responses to the sunk costs of food-handling time (Figure 3B). Given the greater efficacy of washing (Figure S1) and time needed to carry food to the ocean ( x¯ = 22 ± 15 s; range: 5 to 78 s), there is little incentive to over-wash food (Figure 3B, region I). At the same time, the lowest- and highest-ranking monkeys abstained from washing altogether, choosing instead to minimize food-handling time by over-brushing their food (Figure 3B, region II). This tolerance for fast-diminishing returns underscores the monkeys’ strong aversion to sand; but even so, the long-term benefits of mitigating tooth wear must be balanced against urgent energetic requirements.

Predicted and observed cleaning times.

(A) Mean predicted time (large filled points vs. observed times (violin plots) for brushing and washing food (note log scale). The vertical line associated with predicted times represents the 5 to 95% confidence interval. (B) Predicted cleaning time as a function of cleaning inefficiency c, and handling time h, with mean predicted values (black points) for brushing and washing based on observed cleaning inefficiencies and handling times. The colored points (as in panel A) represent observed cleaning times. The trade-off between longer food handling times and efficient cleaning (oceanside food-washing; Region I) and shorter handling times and inefficient cleaning (immediate food-brushing; Region II) is depicted by the black curve.

Added in the revised version: Discussion

Added in the revised version: Our experiment reveals individual variation at a moment in time. It is a cross-sectional result that can only hint at the long-term fitness consequences of differential cleaning behaviors, a topic with implications for diverse fields of inquiry, from evolutionary theory to paleoanthropology. So, the present discussion was written to put our results into conversation with this wider literature and set an agenda for future research.

Disposable soma

Earlier version: The disposable-soma hypothesis of senescence predicts investment in the immediate needs of survival or reproduction over tooth preservation (Carranza et al., 2004). Dominant monkeys face this predicament because rapid food intake rates are integral to sustaining dominance and accruing reproductive success. For dominant males, energy intake determines their ability to sustain consortships during mating (Higham et al., 2011); and, for dominant females, it affects practically every measure of fitness (Alberts, 2019; Cooper et al., 2022). Our findings suggest that dominant monkeys refrained from washing to maximize short-term energy intake (Figure 2D). In short, they prioritized pressing energetic needs over the long-term benefits of tooth preservation—a ‘live fast, die young’ life-history strategy. This view of teeth as disposable soma may explain why dominant male monkeys experience faster senescence and earlier mortality (Anderson et al., 2021). Estimating fitness consequences is beyond the scope of our study, but our findings suggest that a prolonged life is also subject to diminishing returns.

Revised version: The disposable-soma hypothesis of senescence predicts investment in the immediate needs of survival or reproduction over tooth preservation (Carranza et al., 2004). It is a familiar predicament for high-ranking monkeys. In longitudinal studies of macaques and baboons, it is evident that social dominance incurs high energy costs that are met with high feeding rates (Higham et al., 2011; Gesquiere et al., 2025), behaviors that contribute positively to female fitness (van Noordwijk and van Schaik, 1999; Alberts, 2019; Cooper et al., 2022). Our findings raise the possibility that dominant monkeys refrained from washing to maximize feeding rate (Figure 2D). In other words, they prioritized pressing energetic needs over the long-term benefits of tooth preservation—a ‘live fast, die young’ life-history strategy. This view of teeth as disposable soma could explain why dominant males experience faster senescence and earlier mortality (Anderson et al., 2021). Perhaps even a prolonged life is subject to diminishing returns.

Paleo matter(s)

Earlier version: The full extent of sand-mediated tooth wear is unknown for our study population, but it is probably extreme among the highest-ranked individuals. If affirmed, the findings could affect our views of the hominin fossil record by challenging the assumption that dietary variability is the principal cause of variable dental wear. Some species, notably Paranthropus boisei, had ready access to water, which raises the possibility that they—like many primate species—assiduously washed their food, an essential behavior if their diet featured gritty underground plant tissues (Wrangham et al., 2009; Fannin et al., 2021). Other species, notably P. robustus, have extremely variable levels of tooth pitting (Peterson et al., 2018), which could reflect, at least partially, the absence of extensive wetlands (Herries et al., 2010) coupled with interindividual variation in food-cleaning behaviors. Tellingly, the dental wear observed on the macaques of Koshima Island bears striking similarities to the hominin fossil record (Towle et al., 2022), suggesting that populations of food-cleaning monkeys are a valuable model system that warrant further study.

Revised version: The full extent of sand-mediated tooth wear is unknown for our study population, but it is probably extreme among the highest-ranked individuals. If affirmed, the findings could affect our views of the hominin fossil record by challenging the assumption that dietary variability is the principal cause of variable dental wear. Some species, notably Paranthropus boisei, had ready access to water, which raises the possibility that they—like many primate species—assiduously washed their food, an essential behavior if their diet featured gritty underground plant tissues (Wrangham et al., 2009; Fannin et al., 2021). Other species, notably P. robustus, have extremely variable levels of tooth pitting (Peterson et al., 2018), which could reflect, at least partially, the absence of extensive wetlands (Herries et al., 2010) coupled with rank effects on food-cleaning behaviors. Tellingly, the dental wear observed on the macaques of Koshima Island bears striking similarities to the hominin fossil record (Towle et al., 2022), suggesting that populations of food-cleaning monkeys are a valuable model system that warrant further study.

Significance

Earlier version: Our study leverages a new method in ecological research to provide the first analysis of siliceous particles on primate foods. Our experiment probes the behavioral economics of wild monkeys, revealing a strong aversion to sandy foods. Yet, the monkeys behaved irrationally when cleaning their foods, allocating excess time than predicted by an optimization model. Some individuals fell victim to the sunk cost fallacy by over-washing their foods (Box 2), whereas dominant monkeys abstained from washing altogether, seemingly sacrificing their teeth at the altar of high rank, a social status that relies rapid food intake. Our results support the disposable-soma hypothesis for senescence while kicking the tires of a treasured assumption in paleoanthropology.

Revised version: Our study leverages a new method in ecological research to provide the first analysis of siliceous particles on primate foods. Our experiment investigates the behavioral economics of wild monkeys, revealing a strong aversion to sand in their mouths. Yet, the monkeys behaved irrationally when cleaning their foods, allocating excess time than predicted by an optimization model. Some individuals fell victim to the sunk cost fallacy by over-washing their foods (Box 2), whereas dominant monkeys abstained from washing altogether, seemingly sacrificing their teeth on the altar of high rank, a social status that depends on rapid food intake. Our results are compatible with the disposable-soma hypothesis for senescence, and they call into question some treasured assumptions in paleoanthropology.

Earlier version: Box 2. Balancing sunk costs against future benefits

Revised version: Balancing sunk costs against future benefits

The Aérospatiale-BAC Concorde 102 is an iconic aircraft manufactured from 1965–1979. This photo shows British Airways flight 002 on the eve of its final commercial flight on October 24, 2003 (photograph by Richard Vandervord, reproduced with permission).

Carrying food to the ocean costs time and energy. Some monkeys stood upright and walked to the water because their hands held cucumbers, paying an extremely high energetic cost (Nakatsukasa et al., 2006). Such irrecoverable expenses—or ‘sunk costs’—should not sway the optimal washing time of rational monkeys, which we calculated as 2.40 ± 0.74 sec per cucumber slice. Still, many monkeys washed their food far beyond the point of diminishing returns (Figure 3A). This level of over-washing speaks to their aversion to sand, but it also suggests that some individuals succumbed to the sunk cost fallacy, or “Concorde Effect” (Arkes and Ayton, 1999).

The Concorde is a supersonic aircraft. Production began in 1962 as a joint Anglo-French enterprise, but significant cost overruns undermined any chance of profit. By 1971, the opening line of a British central policy review staff memorandum was explicit and brutal: “Concorde is a commercial disaster,” it said. “It should never have been started.” Still, both governments continued to pour millions into the project on the grounds that they had already placed substantial investment in it. Such reasoning is viewed as an economic fallacy because decisions should be based on prospective (future) costs only. The same memo put it this way, “The decision whether or not to abandon Concorde must start from where we are now—much of the milk is already spilt.”

Irrational decisions are not unique to humans. When faced with a foraging task that required inaction, humans, rats, and mice each fell victim to the Concorde Effect; they were more likely to complete a trial (i.e., to continue waiting instead of opting out) the longer they had already waited (Sweis et al., 2018). In another experiment, Watzek and Brosnan (2020) showed that tufted capuchins and rhesus macaques are reluctant to forfeit a small investment when completing a psychomotor task, persisting 5–7 times longer than was optimal for a food reward. A criticism of these findings is that animal hunger can bias the results, impelling irrational decisions (Ott et al., 2022). In our experiment, however, middle-ranking monkeys delayed consumption of a food reward already in their possession, suggesting an incentive to offset the long-terms costs of sand-mediated tooth damage. If irrational over-washing confers adaptive benefits by prolonging the functional life of teeth, then there could be a kernel wisdom in the Concorde Effect.

Methods and materials

Study site and population

Earlier version: Koram Island (12.242°, 100.009°) lies ~1 km offshore in the Gulf of Thailand and within Khao Sam Roi Yot National Park, Prachuap Khiri Khan, Thailand. It has an area of 0.45 km2 and a coastline of 3.5 km. The habitat—limestone karst blanketed with a dense flora of dwarf evergreen trees and deciduous scrub, and encircled by rocky shore and sandy beaches—supports a population of ca. 75 long-tailed macaques described as hybrids at the subspecies taxonomic level (Macaca fascicularis aurea × M.f. fascicularis)(Gumert et al., 2019). The animals are well habituated to human observers due to regular tourism and sustained study since 2013 (Tan et al., 2018). Most of this research has revolved around stone tool-mediated foraging on mollusks, the only activity known to elicit stone handling (Malaivijitnond et al., 2007; Gumert and Malaivijitnond, 2012, 2013; Tan et al., 2015), although infants and juveniles will sometimes use stones during object play (Tan, 2017). There has been no prior examination of food-cleaning behaviors.

Revised version: Koram Island (12.242°, 100.009°) lies ~1 km offshore in the Gulf of Thailand and within Khao Sam Roi Yot National Park, Prachuap Khiri Khan, Thailand. It has an area of 0.45 km2 and a coastline of 3.5 km. The habitat—limestone karst blanketed with a dense flora of dwarf evergreen trees and deciduous scrub, and encircled by rocky shore and sandy beaches—supports a population of ca. 75 long-tailed macaques described as hybrids at the subspecies taxonomic level (Macaca fascicularis aurea x M.f. fascicularis) (Gumert et al.,2019). The animals are well habituated to human observers due to regular tourism and sustained study since 2013 (Tan et al., 2018). Most of this research has revolved around stone tool-mediated foraging on mollusks, the only activity known to elicit stone handling (Malaivijitnond et al., 2007; Gumert and Malaivijitnond, 2012, 2013; Tan et al., 2015), although infants and juveniles will sometimes use stones during object play (Tan, 2017). There has been no prior examination of food-cleaning behaviors.

Rank determination

Added in the revised version: Macaques formmulti-malemulti-female (polygynandrous) social groupswith individual dominance hierarchies. In M. fascicularis, the hierarchy is strictly linear and extremely steep, meaning aggression is unidirectional (de Waal, 1977; van Noordwijk and van Schaik, 2001) with profound asymmetries in outcomes for individuals of adjacent ranks (Balasubramaniam et al., 2012). Further, the dominance hierarchies of philopatric females are stable and predictable. Daughters follow the pattern of youngest ascendancy, ranking just below their mothers with few known exceptions among older sisters (de Waal, 1977; van Noordwijk and van Schaik, 1999). Taken together, these species traits are conducive to unequivocal rank determinations.

Added in the revised version: To determine the rank-order of adults in our study group, we recorded dyadic agonistic interactions and their outcomes (i.e., aggression, supplants, and silent-bared-teeth displays of submission) during 5-min focal follows of individuals based on a randomized order of continuous rotation (Tan et al., 2018). These data were supplemented with ad libitum observations and all rank determinations were updated monthly, and when males immigrated or emigrated. This protocol predates our experiment in July-August 2018, representing ∼970 hr of focal data during five years of systematic study (2013-2018).

Earlier version: Macaques form multi-male multi-female (polygynandrous) social groups with individual dominance hierarchies. Among females, this hierarchy is strictly linear and stable through time (van Noordwijk and van Schaik, 1999). To determine the rank-order of adults, we recorded dyadic agonistic interactions and their outcomes (i.e., aggression, supplants, and silent-bared-teeth displays of submission) during 5-min focal follows of individuals based on a randomized order of continuous rotation (Tan et al., 2018). In some cases, these data were supplemented with ad libitum observations. This protocol existed during five years (2013–2018) of continual observations before we conducted our experiment in July-August 2018. To determine the effects of dominance rank on individual food-cleaning propensities, we followed the methods of Levy et al. (2020). We calculated ordinal ranks between 1 (highest) and n (lowest), where n is the number of animals aged ≥ 5 years. Then we combined males (n = 8) and females (n = 16) into a single standardized ordinal ranking of 24 animals. This method of rank determination is well suited for conditions involving densitydependent competition (Levy et al., 2020), such as those of our experiment and more generally on Koram Island, where preferred food resources are limited (Luncz et al., 2017).

Revised version: To determine the effects of dominance rank on individual food-cleaning propensities, we followed the methods of Levy et al. (2020). We calculated ordinal ranks between 1 (highest) and n (lowest), where n is the number of animals aged ≥ 5 years. Then we combined males (n = 8) and females (n = 16) into a single standardized ordinal ranking of 24 animals. This method of rank determination is well suited for conditions involving density-dependent competition (Levy et al., 2020), such as those of our experiment and more generally on Koram Island, where preferred food resources are limited (Luncz et al., 2017).

Measuring sand

To quantify the amount of sand on food surfaces, whether provisioned by tourists (cucumbers, melon, pineapple) or used in our experiment (sliced cucumbers), we applied a quick-drying liquid polymer—granulated plastic (Pioloform BL 16; Wacker-Chemie GMBH, Munich, Germany) mixed with ethanol (18% plastic to 82% ethanol, by weight)—to each food item. When dried, we peeled and stored the sand-infused film for analysis. The advantages of this method are twofold: the removal of exogenous particulate matter is extremely thorough; and the plastic does not detach biogenic silica such as trichromes or phytoliths (Hinton et al., 1996). In the lab, we dissolved each peel in ethanol and separated the sand by centrifugation, producing a pellet. We dried and weighed the pellet, dividing the mass by the surface area of the food object, which we calculated from digital photographs imported into ImageJ v. 1.52. This method produces an estimate of exogenous particulate mass per area (mg mm-2), allowing direct comparison of apples and oranges.

Earlier version: To measure the elemental and physical properties of sand, we dispersed and filtered the pellets in water using a 0.2-μm isopore membrane filter, which we submitted for scanning electron microscopy (SEM) and energy dispersive X-ray spectroscopy (EDS). To establish the parameters for multi-field and bulk analysis, we imaged a representative area of the filter at multiple magnifications and performed discrete particle analysis. 50× magnification allowed for statistically significant representation of particle number and size range (allowing a 5 μm lower particle size range in analysis). All discrete particle analyses indicated silicon rich particles and composition distribution bins were established to include dominant accompanying elements. After establishing these parameters, we initiated multi-field automated analysis using six fields ofviewof the debris field (at 50×). A composition classification was assigned to each particle and data sorted by composition classification and particle size (particle sizing was binned using standard Feret maximum parameter). The sizing bins are standard ISO-16232 size classes.

Revised version: To measure the elemental and physical properties of sand, we dispersed and filtered the pellets in water using a 0.2-μιτι isopore membrane filter, which we submitted for scanning electron microscopy (SEM) and energy dispersive X-ray spectroscopy (EDS). To establish the parameters for multi-field and bulk analysis, we imaged a representative area of the filter at multiple magnifications and performed discrete particle analysis. 50χ magnification allowed for statistically significant representation of particle number and size range (allowing a 5 μm lower particle size range in analysis). All discrete particle analyses indicated silicon rich particles and composition distribution bins were established to include dominant accompanying elements. After establishing these parameters, we initiated multi-field automated analysis using six fields of view of the debris field (at 50χ). A composition classification was assigned to each particle and data sorted by composition classification and particle size (particle sizing was binned using standard Feret maximum parameter). The sizing bins are standard ISO-16232 size classes.

Experimental design

To elicit food handling behaviors and determine individual cleaning preferences, we put three cucumber slices in each of three trays (20 × 30 × 10 cm) and manipulated the amount of contaminating sand. In the low-sand treatment, we put cucumber slices in a tray without sand; however, contamination from aeolian sand was unavoidable. In the intermediate-sand treatment, we lined the tray with sand and put the cucumber slices on the surface. In the high-sand treatment, we buried cucumber slices completely (Figure 2A). Trays were placed 1.5 m apart and 15 m from the ocean, and we randomized the color and sequence of trays across trials.

Earlier version: Trials began when one or more monkeys approached the trays and ended when the animals finished every cucumber slice or abandoned the experiment (range: 10 s to 14 min). We used video recordings to determine the onset and offset of individual food-handling bouts, beginning from initial contact with a cucumber slice and ending when the final slice, in its entirety, entered the mouth. Within each bout, we determined the duration of brushing and washing behaviors, defining each from the onset of serial stereotypical forelimb movements to the moment of oral ingestion. We estimated energy intake rates by calculating the number of cucumber slices consumed during each food-handling bout, and multiplying each slice by 1.1 kcal (source: U.S. Department of Agriculture, FoodData Central, 2019) and 4.186 kJ (Hargrove, 2007). We performed 101 trials over 5 weeks and recorded 1,282 food-handling bouts by 42 individual monkeys. To minimize the potential confounding effects of dominance interactions, we analyzed trials with ≤ 3 monkeys. Thus, 937 food-handling bouts were included in GLMM statistical models, which included data on individual rank, sex, and sand treatment. Ifa monkey consumed a cucumber slice without brushing or washing it, the zero-second duration was included in both GLMMs.

Revised version: Trials began when one or more monkeys approached the trays and ended when the animals finished every cucumber slice or abandoned the experiment (range: 10 s to 14 min). We used video recordings to determine the onset and offset of individual food-handling bouts, beginning from initial contact with a cucumber slice and ending when the final slice, in its entirety, entered the mouth. Within each bout, we determined the duration of brushing and washing behaviors, defining each from the onset of serial stereotypical forelimb movements to themoment of oral ingestion. We estimated energy intake rates by calculating the number of cucumber slices consumed during each food-handling bout, and multiplying each slice by 1.1 kcal (source: U.S. Department of Agriculture, FoodData Central, 2019) and 4.186 kJ (Hargrove, 2007). We performed 101 trials over 5 weeks and recorded 1,282 food-handling bouts by 42 individual monkeys. We excluded trials from analysis if the number of participating monkeys exceeded the number of feeding stations, as these conditions produced high levels of feeding competition with scant cleaning behavior. Such conditions effectively erased individual variation in sand removal, the topic motivating our experiment. Accordingly, we analyzed trials with ≤ 3 monkeys, putting 937 food-handling bouts into the GLMM statistical models, which included data on individual rank, sex, and sand treatment. If a monkey consumed a cucumber slice without brushing or washing it, the zero-second duration was included in both GLMMs.

Behavioral analyses

To model variance in brushing and washing behaviors as a function of experimental treatment and rank, we fit generalized linear mixed models (GLMMs) using the glmmTMB package in R version 4.4.1 (Brooks et al., 2017). We then evaluated the performance of each model using the standardized residual diagnostic tools for hierarchical regression models available in the DHARMa package (Hartig, 2022). In each GLMM, we modelled sand treatment (a categorical variable with three levels), sex (a categorical variable with two levels), and ordinal rank (a discrete variable ranging from 1 to 24) as fixed effects. For our brushing model, we incorporated an additional interaction term as a fixed effect: sand treatment × ordinal rank. For our washing dataset, we ran two models: the first modeled the relationship between ordinal rank and food cleaning as a linear term, while the second modeled the relationship as a quadratic term. We also incorporated both the linear and quadratic terms into the interaction effect with sand treatment in each washing model. To account for experimental variance among individuals and control for pseudoreplication (because the number of feeding bouts per individual varied widely; Bolker et al. (2009)), we included individual ID as a random intercept.

Our brushing and washing data sets were whole-number counts (seconds) with means < 5. The distributions were right-skewed with high concentrations of biologically-meaningful zeros (Martin et al., 2005) (i.e., instances of food-handling without any cleaning behavior). Thus, we fit a series of zero-inflated generalized linear mixed models (ZIGLMM), each with a logit-link function, a single zero-inflation parameter applying to all observations, and Poisson error distribution. For the foodbrushing model, we fit a zero-inflated Poisson (ZIP), which produced favorable standardized residual diagnostic plots with no major patterns of deviation (Figure S3) and minor, but non-significant, underdispersion (DHARMa dispersion statistic = 0.99, p = 0.80). For our two food-washing models, we used zero-inflated models with Conway-Maxwell Poisson (ZICMP) distributions, an error distribution chosen for its ability to handle data that is more underdispersed (DHARMa dispersion statistic = 8.2E-09, p = 0.74) than the standard zero-inflated Poisson (Brooks et al., 2019). Using this error distribution improved residual diagnostic plots over a standard ZIP model and we view any deviations in the standardized residuals as minor and due to the smaller sample size of our food-washing data set (Figures S4 and S5) (Hartig, 2022). We report the summarized fixed effects tests for each GLMM in Tables S1–S3 as Analysis of Deviance Tables (Type II Wald chi square tests, one-sided) along with χ2 values, degrees of freedom, and p-values (one-sided tests). Full model summaries with standard errors and confidence intervals are also included in Tables S4–S6. For all statistical analyses, we set α = 0.05.

Optimal cleaning time model

Earlier version: To model the optimization of sand removal, we drew inspiration from the marginal value theorem of Charnov (1976), defining two temporal periods: handling time h, which includes an assessment time and pre-cleaning time, and cleaning time t. Assessment time (set as a constant 1 second) includes visual fixation on a food object and forelimb extension before contact, whereas pre-cleaning time represents all handling activities that precede cleaning. During brushing, the pre-cleaning time was essentially nil (zero seconds), but washing required travel from the experimental treatments to the ocean, requiring longer pre-cleaning times ( x¯ = 22 ± 15 s; range: 5 to 78 s). We assumed that the proportion of sand removed from each cucumber follows the saturating relationship g(t) = t/(c +t), where c is the cleaning inefficiency, or the half-saturation constant associated with brushing or washing. As c increases, so does the inefficiency of a given cleaning behavior. Given our observations, it requires on average 2.97 s of brushing to remove 75±7% of grit, and 3.53 s of washing to remove 93 ± 4% of grit (Figure S1). From these data, and including the experimental uncertainty associated with grit removal percentages, we obtained distributions for estimated cleaning inefficiencies, cbrushing = 0.99 ± 0.38 s and cwashing = 0.27 ± 0.15 s, such that washing (without considering handling costs) is the most efficient strategy. The rate of grit removal is then given by R(t) = g(t)/(h +t), which reaches a maximum at the optimal cleaning time t*=ch . For brushing and washing cleaning strategies, we obtain the expected optimal cleaning times tbrushing* = 0.98±0.19 s, and twashing* = 2.40 ± 0.74 s (Figure 3a), respectively, and where including additional sources of uncertainty did not alter our findings (Figure S6). These optimal cleaning times are defined exclusively with respect to maximizing the rate of grit removal, without considering the potentially cascading effects of these strategies on fitness.

Revised version: To model the optimization of sand removal, we drew inspiration from the marginal value theorem of Charnov (1976), defining two temporal periods: handling time h, which includes an assessment time and pre-cleaning time, and cleaning time t. Assessment time (set as a constant 1 second) includes visual fixation on a food object and forelimb extension before contact, whereas pre-cleaning time represents all handling activities that precede cleaning. During brushing, the pre-cleaning time was essentially nil (zero seconds), but washing required travel from the experimental treatments to the ocean, requiring longer pre-cleaning times ( x¯ = 22 ± 15 s; range: 5 to 78 s). We assumed that the proportion of sand removed from each cucumber follows the saturating relationship g(t) = t/(c +1), where c is the cleaning inefficiency, or the half-saturation constant associated with brushing or washing. As c increases, so does the inefficiency of a given cleaning behavior. Given our observations, it requires on average 2.97 s of brushing to remove 75±7% of grit, and 3.53 s of washing to remove 93 ± 4% of grit (Figure S1). From these data, and including the experimental uncertainty associated with grit removal percentages, we obtained distributions for estimated cleaning inefficiencies, cbrushing = 0.99 ± 0.38 s and cwashing = 0.27 ± 0.15 s, such that washing (without considering handling costs) is the most efficient strategy. The rate of grit removal is then given by R(t) = g(t)/(h + 1), which reaches a maximum at the optimal cleaning time t*=ch . For brushing and washing cleaning strategies, we obtain the expected optimal cleaning times tbrushing*=0.98±0.19 s, and twashing*=2.40±0.74 s (Figure 3a), respectively, and where including additional sources of uncertainty did not alter our findings (Figure S6). These optimal cleaning times are defined exclusively with respect to maximizing the rate of grit removal, without considering the potentially cascading effects of these strategies on fitness.